---
title: Exercise 1
description: >-
  Exercise 1


---
## Exercise 1

```yaml
type: NormalExercise
lang: r
xp: 100
skills: 1
key: 2bdda055ef
```

Do some data science.

`@instructions`


`@hint`


`@pre_exercise_code`
```{r}
fgdfgdfg
```
`@sample_code`
```{r}
dfgdfgdfgdfg
```
`@solution`
```{r}
dfgdfgdfg
```
`@sct`
```{r}
dfgdfgdfgfdg
```




